unsat
(declare-const prt Bool)
(declare-const red_SR Bool)
(declare-const org_MR Bool)
(declare-const grn Bool)
(declare-const org Bool)
(declare-const rd1 Bool)
(declare-const grn_SR Bool)
(declare-const org_SR Bool)
(declare-const red_MR Bool)
(declare-const rd2 Bool)
(declare-const circuit Bool)
(define @t1 () (not rd2))
(define @t2 () (not rd1))
(define @t3 () (not org))
(define @t4 () (and @t2 @t1))
(define @t5 () (and prt rd1))
(define @t6 () (or org @t5))
(define @t7 () (not @t5))
(define @t8 () (not @t4))
(define @t9 () (not @t6))
(assume @p1 circuit)
(assume @p2 rd2)
(assume @p3 red_MR)
(assume @p4 (or grn_SR org_SR))
(assume @p5 (=> grn (and @t3 @t2 @t1)))
(assume @p6 (=> org @t4))
(assume @p7 (=> rd1 @t1))
(assume @p8 (= org_MR org))
(assume @p9 (= grn_SR rd1))
(assume @p10 (= org_SR rd2))
(assume @p11 (= red_SR (or grn org)))
(assume @p12 (not (= rd1 @t6)))
(step @p13 (or @t2 @t1) :rule implies_elim :premises (@p7))
(step @p14 @t2 :rule resolution :premises (@p13 @p2) :args (false rd2))
(step @p15 (or @t7 rd1) :rule cnf_and_pos :args (@t5 1))
(step @p16 @t7 :rule resolution :premises (@p15 @p14) :args (true rd1))
(step @p17 (or @t8 @t1) :rule cnf_and_pos :args (@t4 1))
(step @p18 @t8 :rule resolution :premises (@p17 @p2) :args (false rd2))
(step @p19 (or @t3 @t4) :rule implies_elim :premises (@p6))
(step @p20 @t3 :rule resolution :premises (@p19 @p18) :args (true @t4))
(step @p21 (or @t9 org @t5) :rule cnf_or_pos :args (@t6))
(step @p22 @t9 :rule chain_resolution :premises (@p21 @p20 @p16) :args ((@list true true) (@list org @t5)))
(step @p23 (or rd1 @t6) :rule not_equiv_elim1 :premises (@p12))
(step @p24 @t6 :rule resolution :premises (@p23 @p14) :args (true rd1))
(step @p25 false :rule contra :premises (@p24 @p22))

