(declare-const prt Bool)
(declare-const red_MR Bool)
(declare-const circuit Bool)
(declare-const red_SR Bool)
(declare-const org_SR Bool)
(declare-const grn_SR Bool)
(declare-const org_MR Bool)
(declare-const grn_MR Bool)
(declare-const grn Bool)
(declare-const org Bool)
(declare-const rd1 Bool)
(declare-const rd2 Bool)
(define @t1 () (not rd2))
(define @t2 () (not rd1))
(define @t3 () (not org))
(define @t4 () (and @t3 @t2 @t1))
(define @t5 () (and @t2 @t1))
(define @t6 () (and prt grn))
(define @t7 () (= org @t6))
(define @t8 () (not @t4))
(define @t9 () (not grn))
(define @t10 () (not @t6))
(define @t11 () (not @t5))
(assume @p1 (=> grn @t4))
(assume @p2 (=> org @t5))
(assume @p3 (=> rd1 @t1))
(assume @p4 (= grn_MR grn))
(assume @p5 (= org_MR org))
(assume @p6 (= grn_SR rd1))
(assume @p7 (= org_SR rd2))
(assume @p8 (= red_SR (or grn org)))
(assume @p9 (= red_SR (or grn_MR org_MR)))
(assume @p10 circuit)
(assume @p11 rd2)
(assume @p12 red_MR)
(assume @p13 (or grn_SR org_SR))
(assume @p14 (not @t7))
(step @p15 (or @t8 @t1) :rule cnf_and_pos :args (@t4 2))
(step @p16 @t8 :rule resolution :premises (@p15 @p11) :args (false rd2))
(step @p17 (or @t9 @t4) :rule implies_elim :premises (@p1))
(step @p18 @t9 :rule resolution :premises (@p17 @p16) :args (true @t4))
(step @p19 (or @t10 grn) :rule cnf_and_pos :args (@t6 1))
(step @p20 @t10 :rule resolution :premises (@p19 @p18) :args (true grn))
(step @p21 (or @t11 @t1) :rule cnf_and_pos :args (@t5 1))
(step @p22 @t11 :rule resolution :premises (@p21 @p11) :args (false rd2))
(step @p23 (or @t3 @t5) :rule implies_elim :premises (@p2))
(step @p24 @t3 :rule resolution :premises (@p23 @p22) :args (true @t5))
(step @p25 (or @t7 org @t6) :rule cnf_equiv_neg1 :args (@t7))
(step @p26 @t7 :rule chain_resolution :premises (@p25 @p24 @p20) :args ((@list true true) (@list org @t6)))
(step @p27 false :rule contra :premises (@p26 @p14))
